# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 3-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np
#pd.options.display.float_format = '{:.5f}'.format
# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

```python
!pip install xlrd
```

Output:
```
Requirement already satisfied: xlrd in /opt/conda/lib/python3.7/site-packages (1.2.0)
```

```python
import xlrd
Energy = pd.read_excel("assets/Energy Indicators.xls")
Energy.columns
```

Output:
```
Index(['Unnamed: 0', 'Unnamed: 1', 'Unnamed: 2', 'Unnamed: 3', 'Unnamed: 4',
       'Unnamed: 5'],
      dtype='object')
```

```python
Energy.head(20)
```

Output:
```
Unnamed: 0   Unnamed: 1  \
0          NaN          NaN   
1          NaN          NaN   
2          NaN          NaN   
3          NaN          NaN   
4          NaN          NaN   
5          NaN          NaN   
6          NaN          NaN   
7          NaN          NaN   
8          NaN          NaN   
9          NaN          NaN   
10         NaN          NaN   
11         NaN          NaN   
12         NaN          NaN   
13         NaN          NaN   
14         NaN          NaN   
15         NaN          NaN   
16         NaN          NaN   
17         NaN  Afghanistan   
18         NaN      Albania   
19         NaN      Algeria   

                                           Unnamed: 2     Unnamed: 3  \
0                    Environmental Indicators: Energy            NaN   
1                                                 NaN            NaN   
2   Energy Supply and Renewable Electricity Produc...            NaN   
3                                                 NaN           
```

```python
Energy.drop([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16], inplace=True)
Energy.head(5)
```

Output:
```
Unnamed: 0      Unnamed: 1      Unnamed: 2 Unnamed: 3 Unnamed: 4  \
17         NaN     Afghanistan     Afghanistan        321         10   
18         NaN         Albania         Albania        102         35   
19         NaN         Algeria         Algeria       1959         51   
20         NaN  American Samoa  American Samoa        ...        ...   
21         NaN         Andorra         Andorra          9        121   

   Unnamed: 5  
17    78.6693  
18        100  
19    0.55101  
20   0.641026  
21    88.6957
```

```python
#The first two columns are unneccessary
Energy.drop(columns=["Unnamed: 0", "Unnamed: 1"], inplace=True)
```

```python
#change the column labels to ['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]
Energy.rename(columns={"Unnamed: 2":'Country',
                       "Unnamed: 3":'Energy Supply',
                       "Unnamed: 4":'Energy Supply per Capita',
                       "Unnamed: 5":'% Renewable'}, inplace=True)
Energy.head()
```

Output:
```
Country Energy Supply Energy Supply per Capita % Renewable
17     Afghanistan           321                       10     78.6693
18         Albania           102                       35         100
19         Algeria          1959                       51     0.55101
20  American Samoa           ...                      ...    0.641026
21         Andorra             9                      121     88.6957
```

```python
#For all countries which have missing data (e.g. data with "...") make sure this is reflected as np.NaN values.
Energy.replace('...',np.NaN, inplace=True)
Energy.head()
```

Output:
```
Country  Energy Supply  Energy Supply per Capita  % Renewable
17     Afghanistan          321.0                      10.0    78.669280
18         Albania          102.0                      35.0   100.000000
19         Algeria         1959.0                      51.0     0.551010
20  American Samoa            NaN                       NaN     0.641026
21         Andorra            9.0                     121.0    88.695650
```

```python
#Convert Energy Supply to gigajoules (Note: there are 1,000,000 gigajoules in a petajoule).
Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: x*1000000, 1)
Energy.head()
```

Output:
```
Country  Energy Supply  Energy Supply per Capita  % Renewable
17     Afghanistan   3.210000e+08                      10.0    78.669280
18         Albania   1.020000e+08                      35.0   100.000000
19         Algeria   1.959000e+09                      51.0     0.551010
20  American Samoa            NaN                       NaN     0.641026
21         Andorra   9.000000e+06                     121.0    88.695650
```

```python
#Energy['Energy Supply'] = Energy['Energy Supply'].apply(lambda x: '%.1f' % x)
#Energy.head()
```

```python
#'Bolivia (Plurinational State of)' should be 'Bolivia'. 'Switzerland17' should be 'Switzerland'
import re
def make_names_pretty(country):
    return str(re.findall(r"[A-Za-z\-\s]+", str(country))[0])

a = 'Bolivia (Plurinational State of)'
b = 'Switzerland17'
c = 'United states nkl anj'
d = "China, blabla"
print(make_names_pretty(a), make_names_pretty(c), make_names_pretty(d))
```

Output:
```
Bolivia  United states nkl anj China
```

```python
Energy["Country"] = Energy["Country"].apply(make_names_pretty)
Energy.head()
```

Output:
```
Country  Energy Supply  Energy Supply per Capita  % Renewable
17     Afghanistan   3.210000e+08                      10.0    78.669280
18         Albania   1.020000e+08                      35.0   100.000000
19         Algeria   1.959000e+09                      51.0     0.551010
20  American Samoa            NaN                       NaN     0.641026
21         Andorra   9.000000e+06                     121.0    88.695650
```

```python
#"Republic of Korea": "South Korea",
#"United States of America": "United States",
#"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
#"China, Hong Kong Special Administrative Region": "Hong Kong"
Energy.replace({"Republic of Korea": "South Korea",
                "United States of America": "United States",
                "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
                "China, Hong Kong Special Administrative Region": "Hong Kong"}, inplace=True)
```

```python
Energy["Country"] = Energy["Country"].apply(lambda x: x.strip())
```

```python
GDP = pd.read_csv("assets/world_bank.csv")
GDP.drop([0, 1, 2], inplace=True)
GDP.head()
```

Output:
```
Data Source World Development Indicators  \
3  Country Name                 Country Code   
4         Aruba                          ABW   
5       Andorra                          AND   
6   Afghanistan                          AFG   
7        Angola                          AGO   

                                 Unnamed: 2      Unnamed: 3  Unnamed: 4  \
3                            Indicator Name  Indicator Code      1960.0   
4  GDP at market prices (constant 2010 US$)  NY.GDP.MKTP.KD         NaN   
5  GDP at market prices (constant 2010 US$)  NY.GDP.MKTP.KD         NaN   
6  GDP at market prices (constant 2010 US$)  NY.GDP.MKTP.KD         NaN   
7  GDP at market prices (constant 2010 US$)  NY.GDP.MKTP.KD         NaN   

   Unnamed: 5  Unnamed: 6  Unnamed: 7  Unnamed: 8  Unnamed: 9  ...  \
3      1961.0      1962.0      1963.0      1964.0      1965.0  ...   
4         NaN         NaN         NaN         NaN         NaN  ...   
5         NaN         NaN         NaN         NaN     
```

```python
col_names=list(GDP.loc[3])
col_prev_names = list(GDP.columns)

zip_iterator = zip(col_prev_names, col_names)
names_dict = dict(zip_iterator)

GDP.rename(columns = names_dict, inplace=True)
GDP.head()
```

Output:
```
Country Name  Country Code                            Indicator Name  \
3  Country Name  Country Code                            Indicator Name   
4         Aruba           ABW  GDP at market prices (constant 2010 US$)   
5       Andorra           AND  GDP at market prices (constant 2010 US$)   
6   Afghanistan           AFG  GDP at market prices (constant 2010 US$)   
7        Angola           AGO  GDP at market prices (constant 2010 US$)   

   Indicator Code  1960.0  1961.0  1962.0  1963.0  1964.0  1965.0  ...  \
3  Indicator Code  1960.0  1961.0  1962.0  1963.0  1964.0  1965.0  ...   
4  NY.GDP.MKTP.KD     NaN     NaN     NaN     NaN     NaN     NaN  ...   
5  NY.GDP.MKTP.KD     NaN     NaN     NaN     NaN     NaN     NaN  ...   
6  NY.GDP.MKTP.KD     NaN     NaN     NaN     NaN     NaN     NaN  ...   
7  NY.GDP.MKTP.KD     NaN     NaN     NaN     NaN     NaN     NaN  ...   

         2006.0        2007.0        2008.0        2009.0        2010.0  \
3  2.006000e+03  2.007000e+03  2
```

```python
GDP.drop([3], inplace=True)

GDP.replace({"Korea, Rep.": "South Korea", 
            "Iran, Islamic Rep.": "Iran",
            "Hong Kong SAR, China": "Hong Kong"}, inplace=True)
GDP.head()
g = GDP.copy()
```

```python
ScimEn = pd.read_excel("assets/scimagojr-3.xlsx")
ScimEn.head(15)
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      2       United States      96661              94747     792274   
2      3               Japan      30504              30287     223024   
3      4      United Kingdom      20944              20357     206091   
4      5  Russian Federation      18534              18301      34266   
5      6              Canada      17899              17620     215003   
6      7             Germany      17027              16831     140566   
7      8               India      15005              14841     128763   
8      9              France      13153              12973     130632   
9     10         South Korea      11983              11923     114675   
10    11               Italy      10964              10794     111850   
11    12               Spain       9428               9330     123336   
12    13                Iran       8896               8
```

```python
needed_ScimEn = ScimEn[0:15]
needed_GDP = GDP.drop(columns=col_names[1:50])
needed_GDP.rename(columns={"Country Name":"Country"}, inplace=True)
needed_GDP.head()
```

Output:
```
Country        2006.0        2007.0        2008.0        2009.0  \
4        Aruba           NaN           NaN           NaN           NaN   
5      Andorra  4.018196e+09  4.021331e+09  3.675728e+09  3.535389e+09   
6  Afghanistan  1.030523e+10  1.172119e+10  1.214448e+10  1.469733e+10   
7       Angola  5.581103e+10  6.842044e+10  7.787420e+10  7.975320e+10   
8      Albania  9.771760e+09  1.034829e+10  1.112752e+10  1.150029e+10   

         2010.0        2011.0        2012.0        2013.0        2014.0  \
4  2.467704e+09           NaN           NaN           NaN           NaN   
5  3.346317e+09  3.185605e+09  3.129538e+09  3.127550e+09           NaN   
6  1.593680e+10  1.691113e+10  1.935220e+10  1.973134e+10  1.999032e+10   
7  8.247091e+10  8.570262e+10  9.012096e+10  9.626143e+10  1.008863e+11   
8  1.192695e+10  1.223109e+10  1.240477e+10  1.254247e+10  1.279331e+10   

         2015.0  
4           NaN  
5           NaN  
6  2.029415e+10  
7  1.039106e+11  
8  1.312082e+10
```

```python
#Join using the intersection of country names. Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15).
#The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations', 'Citations per document', 'H index', 'Energy Supply', 'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008', '2009', '2010', '2011', '2012', '2013', '2014', '2015'].
df = pd.merge(needed_ScimEn, Energy, how='left', on='Country')
df.head(15)
```

Output:
```
Rank             Country  Documents  Citable documents  Citations  \
0      1               China     127050             126767     597237   
1      1               China     127050             126767     597237   
2      1               China     127050             126767     597237   
3      2       United States      96661              94747     792274   
4      3               Japan      30504              30287     223024   
5      4      United Kingdom      20944              20357     206091   
6      5  Russian Federation      18534              18301      34266   
7      6              Canada      17899              17620     215003   
8      7             Germany      17027              16831     140566   
9      8               India      15005              14841     128763   
10     9              France      13153              12973     130632   
11    10         South Korea      11983              11923     114675   
12    11               Italy      10964              10
```

```python
df = pd.merge(df, needed_GDP, how='left', on='Country')
df.head()
```

Output:
```
Rank        Country  Documents  Citable documents  Citations  \
0     1          China     127050             126767     597237   
1     1          China     127050             126767     597237   
2     1          China     127050             126767     597237   
3     2  United States      96661              94747     792274   
4     3          Japan      30504              30287     223024   

   Self-citations  Citations per document  H index  Energy Supply  \
0          411683                    4.70      138   1.271910e+11   
1          411683                    4.70      138   5.850000e+08   
2          411683                    4.70      138   4.700000e+07   
3          265436                    8.20      230   9.083800e+10   
4           61554                    7.31      134   1.898400e+10   

   Energy Supply per Capita  ...        2006.0        2007.0        2008.0  \
0                      93.0  ...  3.992331e+12  4.559041e+12  4.997775e+12   
1                      82.0  
```

```python
China_ES = float(df.iloc[0]["Energy Supply"]) + float(df.iloc[1]["Energy Supply"]) + float(df.iloc[2]["Energy Supply"])
China_ES
```

Output:
```
127823000000.0
```

```python
#df.at[0, "Energy Supply"] = China_ES
df.drop([1, 2], axis=0, inplace=True)
```

```python
df.set_index("Country", inplace=True)
df.head()
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   

                    Self-citations  Citations per document  H index  \
Country                                                               
China                       411683                    4.70      138   
United States               265436                    8.20      230   
Japan                        61554                    7.31      134   
United Kingdom               37874                    9.84      139   
Russian Federation           12422                    1.85       57   

                    Energy Supply  Ene
```

```python
df_names = list(df.columns)
floats = df_names[-10:]
integers = list(map(int, floats))
integers = list(map(str, integers))
zip_iterator = zip(floats, integers)
names_dict = dict(zip_iterator)

df.rename(columns = names_dict, inplace=True)
gf = df.copy()
df.head()
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   

                    Self-citations  Citations per document  H index  \
Country                                                               
China                       411683                    4.70      138   
United States               265436                    8.20      230   
Japan                        61554                    7.31      134   
United Kingdom               37874                    9.84      139   
Russian Federation           12422                    1.85       57   

                    Energy Supply  Ene
```

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
#This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".
def answer_one():
    return df
    raise NotImplementedError()
```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

```python
GDP.rename(columns={"Country Name":"Country"}, inplace=True)

for i in range(244, 282):
    Energy.drop([i], inplace=True)
```

```python
lf = pd.merge(ScimEn, GDP, how='inner', on='Country')
print("lf", len(lf))
df = pd.merge(ScimEn, Energy, how='inner', on='Country')
print("df", len(df))
af = pd.merge(GDP, Energy, how='inner', on='Country')
print("af", len(af))
summa = len(ScimEn)+  len(GDP) + len(Energy)
print(summa, len(ScimEn),  len(GDP), len(Energy))
print(summa - len(df) - len(af) - len(lf) + 30)
```

Output:
```
lf 166
df 175
af 187
682 191 264 227
184
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    return 156
    raise NotImplementedError()
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

```python
needed_GDP_names = list(needed_GDP.columns)
floats = needed_GDP_names[-10:]
integers = list(map(int, floats))
integers = list(map(str, integers))
zip_iterator = zip(floats, integers)
names_dict = dict(zip_iterator)

needed_GDP.rename(columns = names_dict, inplace=True)
```

```python
ScimEn["Country"] = ScimEn["Country"].apply(lambda x: x.strip())
needed_GDP["Country"] = needed_GDP["Country"].apply(lambda x: x.strip())
```

```python
meaningful = gf[integers].mean(axis=1)

GDP2 = meaningful.sort_values(ascending=False)[:15]
GDP2.head(40)
```

Output:
```
Country
United States         1.536434e+13
China                 6.348609e+12
Japan                 5.542208e+12
Germany               3.493025e+12
France                2.681725e+12
United Kingdom        2.487907e+12
Brazil                2.189794e+12
Italy                 2.120175e+12
India                 1.769297e+12
Canada                1.660647e+12
Russian Federation    1.565459e+12
Spain                 1.418078e+12
Australia             1.164043e+12
South Korea           1.106715e+12
Iran                  4.441558e+11
dtype: float64
```

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->

    
    

```python
def answer_three():
    return GDP2
    raise NotImplementedError()
```

```python
assert type(answer_three()) == pd.Series
```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

```python
years_GDP = gf.loc["United Kingdom", "2015"] - gf.loc["United Kingdom", "2006"]
```

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    return years_GDP
    raise NotImplementedError()
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

```python
gf.head()
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   

                    Self-citations  Citations per document  H index  \
Country                                                               
China                       411683                    4.70      138   
United States               265436                    8.20      230   
Japan                        61554                    7.31      134   
United Kingdom               37874                    9.84      139   
Russian Federation           12422                    1.85       57   

                    Energy Supply  Ene
```

```python
mean_energ = gf["Energy Supply per Capita"].mean()
mean_energ
```

Output:
```
157.6
```

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    return mean_energ
    raise NotImplementedError()
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

```python
max_Ren = gf["% Renewable"].max()
max_Ren
```

Output:
```
69.64803
```

```python
select_index = int(np.where(gf["% Renewable"] == max_Ren)[0])
countr = gf.index[select_index]
```

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    return (countr, max_Ren)
    raise NotImplementedError()
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

```python
gf["Ratio"] = gf["Self-citations"]/gf["Citations"]
max_Ratio = gf["Ratio"].max()

select_index1 = int(np.where(gf["Ratio"] == max_Ratio)[0])
countr1= gf.index[select_index1]
#.argmax()
```

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    return (countr1, max_Ratio)
    raise NotImplementedError()
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

```python
gf["Population"] = gf["Energy Supply"]/gf["Energy Supply per Capita"]
population_rating = gf["Population"].sort_values(ascending=False)[2]

select_index2 = int(np.where(gf["Population"] == population_rating)[0])
country2 = gf.index[select_index2]
country2
```

Output:
```
'United States'
```

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    return country2
    raise NotImplementedError()
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

```python
gf['Citable docs per Capita'] = gf['Citable documents'] / gf["Population"]
```

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
import scipy.stats as stats
import numpy as np
import pandas as pd

def answer_nine():
    # YOUR CODE HERE
    corr, pval=stats.pearsonr(gf['Citable docs per Capita'],gf["Energy Supply per Capita"])
    return corr
    raise NotImplementedError()
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

```python
med = gf["% Renewable"].median()
med
```

Output:
```
17.02028
```

```python
gf["Bool Renewable"] = gf["% Renewable"]
gf["Bool Renewable"] = gf["Bool Renewable"].where(gf["Bool Renewable"] >= med, 0)
gf["Bool Renewable"] = gf["Bool Renewable"].where(gf["Bool Renewable"] < med, 1)
gf["Bool Renewable"]
```

Output:
```
Country
China                 1.0
United States         0.0
Japan                 0.0
United Kingdom        0.0
Russian Federation    1.0
Canada                1.0
Germany               1.0
India                 0.0
France                1.0
South Korea           0.0
Italy                 1.0
Spain                 1.0
Iran                  0.0
Australia             0.0
Brazil                1.0
Name: Bool Renewable, dtype: float64
```

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    return gf["Bool Renewable"]
    raise NotImplementedError()
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

```python
assert list(answer_ten()) == [1, 0, 0, 0, 1, 1, 1, 0, 1, 0, 1, 1, 0, 0, 1]
```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

```python
ContinentDict  = {'China':'Asia', 'United States':'North America', 
                  'Japan':'Asia', 'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 'Canada':'North America', 
                  'Germany':'Europe', 'India':'Asia',
                  'France':'Europe', 'South Korea':'Asia', 
                  'Italy':'Europe', 'Spain':'Europe', 
                  'Iran':'Asia', 'Australia':'Australia', 
                  'Brazil':'South America'}
```

```python
Continents =  ['Asia', 'Australia', 'Europe', 'North America', 'South America']
```

```python
ssmt = pd.DataFrame(columns = ['continent', 'size', 'sum', 'mean', 'std'])
ssmt
```

Output:
```
Empty DataFrame
Columns: [continent, size, sum, mean, std]
Index: []
```

```python
for i in Continents:
    ssmt = ssmt.append({'continent': i, 'size':0, 'sum':0, 'mean':0, 'std':0}, ignore_index=True)
ssmt.set_index("continent", inplace=True)
ssmt
```

Output:
```
size sum mean std
continent                      
Asia             0   0    0   0
Australia        0   0    0   0
Europe           0   0    0   0
North America    0   0    0   0
South America    0   0    0   0
```

```python
for group, frame in gf.groupby(ContinentDict):
    ssmt.loc[group] = [len(frame["Population"]), sum(frame["Population"]), frame["Population"].mean(), frame["Population"].std()]
```

```python
ssmt
```

Output:
```
size          sum         mean          std
continent                                                
Asia             5  2.89867e+09  5.79733e+08  6.79098e+08
Australia        1   2.3316e+07   2.3316e+07          NaN
Europe           6   4.5793e+08  7.63216e+07  3.46477e+07
North America    2  3.52855e+08  1.76428e+08   1.9967e+08
South America    1  2.05915e+08  2.05915e+08          NaN
```

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    return ssmt
    raise NotImplementedError()
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

```python
gf["bins15"] = pd.cut(gf["% Renewable"],5)
```

```python
new_gf = gf.groupby([ContinentDict, gf["bins15"]]).size()
new_gf
```

Output:
```
bins15          
Asia           (2.212, 15.753]     4
               (15.753, 29.227]    1
Australia      (2.212, 15.753]     1
Europe         (2.212, 15.753]     1
               (15.753, 29.227]    3
               (29.227, 42.701]    2
North America  (2.212, 15.753]     1
               (56.174, 69.648]    1
South America  (56.174, 69.648]    1
dtype: int64
```

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    return new_gf
    raise NotImplementedError()

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

```python
gf["Population"]["China"]
```

Output:
```
1367645161.2903225
```

```python
gf["PopEst"] = gf["Population"].apply(lambda x: f"{x:,}")
gf["PopEst"]
```

Output:
```
Country
China                 1,367,645,161.2903225
United States          317,615,384.61538464
Japan                  127,409,395.97315437
United Kingdom         63,870,967.741935484
Russian Federation            143,500,000.0
Canada                  35,239,864.86486486
Germany                 80,369,696.96969697
India                 1,276,730,769.2307692
France                  63,837,349.39759036
South Korea            49,805,429.864253394
Italy                  59,908,256.880733944
Spain                    46,443,396.2264151
Iran                    77,075,630.25210084
Australia              23,316,017.316017315
Brazil                 205,915,254.23728815
Name: PopEst, dtype: object
```

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    return(gf["PopEst"])
    raise NotImplementedError()
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```